```python
# Initialize Otter
import otter
grader = otter.Notebook("pre_finalProject.ipynb")
```

# Final Project: Preliminary Notebook Assignment

At the end of this course, you will have the opportunity to complete a final project to explore a dataset of your choice. 

In this preliminary assignment, the goal will be to find a couple of datasets that are interesting to you. Finding these datasets early will give us time to help you choose which dataset to use for your final project.

**Rules.** You are welcome to discuss your dataset/source with other students, but there is plenty of data out there, so let's try to have unique datasets. We will be working on the final project for several weeks and during the final lab session, you will have a chance to share with the whole class.

**Support.** If you're ever feeling overwhelmed or don't know where to start looking, we are here to help!

To get started, load `datascience`, `numpy`, and `plots`.

```python
# Don't change this cell; just run it. If you need additional libraries for your project, you can add them to this cell.

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
from matplotlib import patches
from ipywidgets import interact, interactive, fixed
import ipywidgets as widgets

```

# Picking a Dataset

Find a dataset that is interesting to you that can be used in the final project.  Specifically, this dataset should have many variables that you think might explain the outcome of one particular variable that will have a binary outcome (or if it has many outcomes, you can collapse the different outcomes into two levels).  

The dataset you pick must meet the following criteria:
1.  One of the variables must be a binary choice.  This is the dependent variable or response variable of interest.
2. Check that the split of observations for the two different levels of the dependent variable are such that the less frequent level has at least 10% of the observations in the dataset.
3. There should be at least 2 categorical variables
4. There should be at least 4 numerical variables
5. There should be at least 100 observations

**Question 1.1** In the cell below:
1. Read in the dataset you chose as a table. Load the table, showing the first 10 rows.
2. Edit the comment to describe where you found this dataset

```python
# Edit this comment to describe where you found this dataset
# Load your dataset into a table

my_data_raw = ...
my_data_raw
```

```python
grader.check("q1_1")
```

**Question 1.2:** In the following cell, describe each of the variables in the dataset.  Are they categorical or numerical? How many observations are there?  Add a code cell below this to show how you found the correct number of observations programmatically.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 1.3:**

What is the dependent or response variable you are interested in?  This is the variable that should have two levels (a binary choice).

For example: (1) Is this movie a thriller or a comedy? (2) Is this amazon order Fraudulent or not? (3) Does this patient have cancer or not? 

In the code cell below, make any necessary adjustments to your data so that the variable is formatted in this way.  Then assign the variable `var` to the column `label` of your table that contains the observations of this variable (**IMPORTANT: the label should be a string, so don't forget the quotes!**).

Make sure that each binary choice represents a minimum of 10% of the observations.

```python
# If you need to make adjustments to your data so that the variable is formatted in 2 levels add the needed code below, before setting var
var = ...
var

```

```python
grader.check("q1_3")
```

**Question 1.4** What do you wonder about the relationship between the other variables and the dependent variable?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Now, we are ready to investigate our data visually!

**Question 1.5:** Think about the numerical variables in the dataset that might be related to each other.  In the cell below, make at least three different scatter plots that show the relationship between different variables while also displaying how each case is classified.

Use the following line of code:

**my_data.scatter(`Column 1`, `Column 2`, group=`label`)**

Replace `Column 1` and `Column 2` with the correct column names of numerical variables(features) you would like to investigate. Replace `label` with the column name of the categorical variable you would like to try to classsify.

Note: The commented code in the cell below is sample code for this.

Look at the scatter plots: do they provide additional insight into your answer for question 1.4? Based on this visual data, feel free to revise your response to question 1.4 if necessary. If you're unsure how to interpret the visualizations, check in with your instructor and ask for help.

```python
# Here is the code placeholder to use, uncomment the line and adjust according to the names if the columns in your dataset:
# my_data.scatter("Column 1", "Column 2", group="label")

...
...
...
```

**Congratulations**: You're DONE with the preliminary notebook assignment for the final project! Nice work.
 
Time to submit.

Tests for q1_1 (points: None)
```
>>> type(my_data_raw) == Table
True
```
```
>>> my_data_raw.num_columns >= 6
True
```
```
>>> my_data_raw.num_rows >= 100
True
```


Tests for q1_3 (points: None)
```
>>> type(my_data_raw.column(var)) == np.ndarray
True
```
```
>>> my_data_raw.group(var).num_rows == 2
True
```
```
>>> my_data_raw.group(var).column('count').item(0) / my_data_raw.num_rows >= 0.1
True
```
```
>>> my_data_raw.group(var).column('count').item(1) / my_data_raw.num_rows >= 0.1
True
```
