```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

# Lab 7: Crime and Penalty

Welcome to Lab 7!

```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)

```

# Lab Warm Up!
We will work together as a class in the following coding cells to prepare you for all sections of this lab. 

**Make sure to come to lab on time so you don't miss points for this warm-up!**

Open the text book and let's go over section [12.1 A/B Testing](https://inferentialthinking.com/chapters/12/1/AB_Testing.html)!

```python
# Run this cell to look at the data on births
births = Table.read_table('baby.csv')
births
```

```python
# Let's look at 2 variables, you can follow along using the example in the book
...
```

Let's discuss a null/alternative hypothesis and test statistic as a class!

*Add lab discussion notes here*

```python
# Just run this cell
def difference_of_means(table, group_label):
    """Takes: name of table,
    column label that indicates the group to which the row belongs
    Returns: Difference of mean birth weights of the two groups"""
    reduced = table.select('Birth Weight', group_label)
    means_table = reduced.group(group_label, np.average)
    means = means_table.column(1)
    return means.item(1) - means.item(0)
```

```python
# Calculating the observed statistic
...
```

Let's get started as a class and complete question 1.1 together!

## 1. A/B Testing

A/B testing is a form of hypothesis testing that allows you to make comparisons between two distributions.

You'll almost never be explicitly asked to perform an A/B test. Make sure you can identify situations where the test is appropriate and know how to correctly implement each step.  

**Question 1.1:** The following statements are the unordered steps of an A/B hypothesis test:

1. Choose a test statistic (typically the difference in means between two categories)

2. Shuffle the labels of the original sample, find your simulated test statistic, and repeat many times

3. Find the value of the observed test statistic

4. Calculate the p-value based off your observed and simulated test statistics

5. Define a null and alternate model

6. Use the p-value and p-value cutoff to draw a conclusion about the null hypothesis

Make an array called `ab_test_order` that contains the correct order of an A/B test, where the first item of the array is the first step of an A/B test and the last item of the array is the last step of an A/B test



```python
ab_test_order = ...
```

```python
grader.check("q1_1")
```

<!-- BEGIN QUESTION -->

**Question 1.2:** If the null hypothesis of an A/B test is correct, should the order of labels affect the differences in means between each group? Why do we shuffle labels in an A/B test? 



_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2: Murder Rates

Punishment for crime has many [philosophical justifications](http://plato.stanford.edu/entries/punishment/#ThePun).  An important one is that fear of punishment may *deter* people from committing crimes.

In the United States, some jurisdictions execute people who are convicted of particularly serious crimes, such as murder.  This punishment is called the *death penalty* or *capital punishment*.  The death penalty is controversial, and deterrence has been one focal point of the debate.  There are other reasons to support or oppose the death penalty, but in this project we'll focus on deterrence.

The key question about deterrence is:

> Through our exploration, does instituting a death penalty for murder actually reduce the number of murders?

You might have a strong intuition in one direction, but the evidence turns out to be surprisingly complex.  Different sides have variously argued that the death penalty has no deterrent effect and that each execution prevents 8 murders, all using statistical arguments!  We'll try to come to our own conclusion.

#### The data

The main data source for this lab comes from a [paper](http://cjlf.org/deathpenalty/DezRubShepDeterFinal.pdf) by three researchers, Dezhbakhsh, Rubin, and Shepherd.  The dataset contains rates of various violent crimes for every year 1960-2003 (44 years) in every US state.  The researchers compiled the data from the FBI's Uniform Crime Reports.

Since crimes are committed by people, not states, we need to account for the number of people in each state when we're looking at state-level data.  Murder rates are calculated as follows:

$$\text{murder rate for state X in year Y} = \frac{\text{number of murders in state X in year Y}}{\text{population in state X in year Y}}*100000$$

(Murder is rare, so we multiply by 100,000 just to avoid dealing with tiny numbers.)

```python
murder_rates = Table.read_table('crime_rates.csv').select('State', 'Year', 'Population', 'Murder Rate')
murder_rates.set_format("Population", NumberFormatter)
```

Murder rates vary over time, and different states exhibit different trends. The rates in some states change dramatically from year to year, while others are quite stable. Let's plot a couple, just to see the variety.

**Question 2.1.** Using the given tables `ak` and `mn` in the code cell below, create a table `ak_mn` with two columns of murder rates (one for Alaska, one for Minnesota), in addition to a column of years. This table will have the following structure:

| Year | Murder rate in Alaska | Murder rate in Minnesota |
|------|-----------------------|--------------------------|
| 1960 | 10.2                  | 1.2                      |
| 1961 | 11.5                  | 1                        |
| 1962 | 4.5                   | 0.9                      |

<center>... (41 rows omitted)</center>



```python
# The next lines are provided for you.  They create a table
# containing only the Alaska information and one containing
# only the Minnesota information.
ak = murder_rates.where('State', 'Alaska').drop('State', 'Population').relabeled(1, 'Murder rate in Alaska')
mn = murder_rates.where('State', 'Minnesota').drop('State', 'Population').relabeled(1, 'Murder rate in Minnesota')

# Fill in this line to make a table like the one pictured above.
ak_mn = ...
ak_mn
```

```python
grader.check("q2_1")
```

**Question 2.2:** Using the table `ak_mn`, draw a line plot that compares the murder rate in Alaska and the murder rate in Minnesota over time. Draw a line plot with years on the horizontal axis and murder rates on the vertical axis. Make sure it includes two lines: one for Alaska murder rates and one for Minnesota murder rates. Create this plot using a single call.



```python
# Draw your line plot here
...
```

Now what about the murder rates of other states? Say, for example, California and New York? Run the cell below to plot the murder rates of different pairs of states.

```python
# Compare the murder rates of any two states by filling in the blanks below

from ipywidgets import interact, interactive, fixed, interact_manual
import ipywidgets as widgets

def state(state1, state2):
    state1_table = murder_rates.where('State', state1).drop('State', 'Population').relabeled(1, 'Murder rate in {}'.format(state1))
    state2_table = murder_rates.where('State', state2).drop('State', 'Population').relabeled(1, 'Murder rate in {}'.format(state2))
    s1_s2 = state1_table.join('Year', state2_table)
    s1_s2.plot('Year')
    plt.show()

states_array = murder_rates.group('State').column('State')

_ = interact(state,
             state1=widgets.Dropdown(options=list(states_array),value='California'),
             state2=widgets.Dropdown(options=list(states_array),value='New York')
            )
```

## 3. The Death Penalty

Some US states have the death penalty, and others don't, and laws have changed over time. In addition to changes in murder rates, we will also consider whether the death penalty was in force in each state and each year.

Using this information, we would like to investigate how the presence of the death penalty affects the murder rate of a state.

**Question 3.1.** We want to know whether the death penalty *causes* a change in the murder rate.  Why is it not sufficient to compare murder rates in places and times when the death penalty was in force with places and times when it wasn't?



_Type your answer here, replacing this text._

<!-- END QUESTION -->

### A Natural Experiment

In order to attempt to investigate the causal relationship between the death penalty and murder rates, we're going to take advantage of a *natural experiment*.  A natural experiment happens when something other than experimental design applies a treatment to one group and not to another (control) group, and we have some hope that the treatment and control groups don't have any other systematic differences.

Our natural experiment is this: in 1972, a Supreme Court decision called *Furman v. Georgia* banned the death penalty throughout the US.  Suddenly, many states went from having the death penalty to not having the death penalty.

As a first step, let's see how murder rates changed before and after the court decision.  We'll define the test as follows:

> **Population:** All the states that had the death penalty before the 1972 abolition.  (There is no control group for the states that already lacked the death penalty in 1972, so we must omit them.)  This includes all US states **except** Alaska, Hawaii, Maine, Michigan, Wisconsin, and Minnesota.

> **Treatment group:** The states in that population, in 1973 (the year after 1972).

> **Control group:** The states in that population, in 1971 (the year before 1972).

> **Null hypothesis:** Murder rates in 1971 and 1973 come from the same distribution.

> **Alternative hypothesis:** Murder rates were higher in 1973 than they were in 1971.

Our alternative hypothesis is related to our suspicion that murder rates increase when the death penalty is eliminated.  

**Question 3.2:** Should we use an A/B test to test these hypotheses? If yes, what is our "A" group and what is our "B" group?



_Type your answer here, replacing this text._

<!-- END QUESTION -->

The `death_penalty` table below describes whether each state allowed the death penalty in 1971.

```python
non_death_penalty_states = make_array('Alaska', 'Hawaii', 'Maine', 'Michigan', 'Wisconsin', 'Minnesota')

def had_death_penalty_in_1971(state):
    """Returns True if the argument is the name of a state that had the death penalty in 1971."""
    # The implementation of this function uses a bit of syntax
    # we haven't seen before.  Just trust that it behaves as its
    # documentation claims.
    return state not in non_death_penalty_states

states = murder_rates.group('State').select('State')
death_penalty = states.with_column('Death Penalty', states.apply(had_death_penalty_in_1971, 0))
death_penalty
```

**Question 3.3:** Use the `death_penalty` and `murder_rates` tables to find murder rates in 1971 for states with the death penalty before the abolition. Create a new table `preban_rates` that contains the same information as `murder_rates`, along with a column `Death Penalty` that contains booleans (`True` or `False`) describing if states had the death penalty in 1971.



```python
# Find the states that had death penalty in 1971
# STEP 1: Combine/join the tables murder_rates and death_penalty
rates_and_penalty = ...
# STEP 2: Reduce the new table rates_penalty to only include the States with the death penalty in the year 1971
preban_rates = ...
preban_rates
```

```python
grader.check("q3_3")
```

**Question 3.4:** Create a table `postban_rates` that contains the same information as `preban_rates`, but for 1973 instead of 1971. `postban_rates` should only contain the states found in `preban_rates`.



```python
# STEP 1: Save the states from table preban_rates in an array states_with_penalty 
states_with_penalty = ...

# STEP 2: start from table murder_rates and use the table method "where()" and the predicate "are.contained_in()" to 
# create a new table for 1973 that only contains the states saved in states_with_penalty
rates_1973 = ...

#STEP 3: Add a column Death Penalty (since this is after the ban, all states should have the penalty set to False)
postban_rates = ...

postban_rates = postban_rates.sort("State")
postban_rates
```

```python
grader.check("q3_4")
```

**Question 3.5:** Combine `preban_rates_copy` and `postban_rates` to create a table `change_in_death_rates` that contains each state's population, murder rate, and whether or not that state had the death penalty for both 1971 and 1973. 


```python
preban_rates_copy = preban_rates.copy()

# HINT: If tbl_1 and tbl_2 are tables, tbl_1.append(tbl_2) will create a new table 
# that includes rows from both tbl_1 and tbl_2. Both tables must have exactly the same columns, in the same order.
change_in_death_rates = ...
change_in_death_rates
```

Run the cell below to view the distribution of death rates during the pre-ban and post-ban time periods.

```python
change_in_death_rates.hist('Murder Rate', group = 'Death Penalty')
```

**Question 3.6:** Create a table `rate_means` that contains the average murder rates for the states that had the death penalty and the states that didn't have the death penalty. It should have two columns: one indicating if the penalty was in place, and one that contains the average murder rate for each group.


```python
# HINTS: Start from change_in_death_rates defined in Question 3.5
# STEP 1: reduce the table to the 2 columns of interest, 
# STEP 2: use the method group() to separate the states with/without Death Penalty 
# and display the average murder rate for each group
rate_means = ...
rate_means
```

```python
grader.check("q3_6")
```

<!-- BEGIN QUESTION -->

**Question 3.7:** We want to figure out if there is a difference between the distribution of death rates in 1971 and 1973. Specifically, we want to test if murder rates were higher in 1973 than they were in 1971. 

What should the test statistic be? How does it help us differentiate whether the data supports the null and alternative? 

If you are in lab, confirm your answer with your instructor or a tutor before moving on.



_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 3.8:** Set `observed_difference` to the observed test statistic using the `rate_means` table 


```python
# STEP 1: Safe the Murder Rate average column of table rate_means into an array 'means'
means = ...
# STEP 2: Calculated the observed test statistic using the array means
observed_difference = ...
observed_difference
```

```python
grader.check("q3_8")
```

**Question 3.9:** Given a table like `change_in_death_rates`, a value column `label`, and a group column `group_label`, write a function that calculates the appropriate test statistic.


```python
# HINT: Combine the steps from question 3.6 and question 3.8, to create a generic function
# Use the input arguments: table, labels_col and values_col
# Make sure you are NOT use change_in_death_rates, "Death Penalty" and "Murder Rate" inside your function definition

def find_test_stat(table, labels_col, values_col):
    ...

find_test_stat(change_in_death_rates, "Death Penalty", "Murder Rate")
```

```python
grader.check("q3_9")
```

When we run a simulation for A/B testing, we resample by shuffling the labels of the original sample. If the null hypothesis is true and the murder rate distributions are the same, we expect that the difference in mean death rates will be not change when "Death Penalty" labels are changed.

**Question 3.10:** Write a function `simulate_and_test_statistic` to compute one trial of our A/B test. Your function should run a simulation and return a test statistic.


```python
# Note: The test here is fairly lenient, if you have an issue with the following questions, make sure to take a look at your answer to 3.9. 
# Specifically, make sure that you are taking the directionality of our alternative hypothesis into account.
def simulate_and_test_statistic(table, labels_col, values_col):
    ...

simulate_and_test_statistic(change_in_death_rates, "Death Penalty", "Murder Rate")
```

```python
grader.check("q3_10")
```

**Question 3.11:** Simulate 5000 trials of our A/B test and store the test statistics in an array called `differences`


```python
# This cell might take a couple seconds to run
differences = make_array()

...
differences
```

```python
grader.check("q3_11")
```

Run the cell below to view a histogram of your simulated test statistics plotted with your observed test statistic

```python
Table().with_column('Difference Between Group Means', differences).hist()
plt.scatter(observed_difference, 0, color='red', s=30, zorder=2);
```

**Question 3.12:** Find the p-value for your test and assign it to `empirical_P`


```python
empirical_P = ...
empirical_P
```

```python
grader.check("q3_12")
```

<!-- BEGIN QUESTION -->

**Question 3.13:** Using a 5% P-value cutoff, draw a conclusion about the null and alternative hypotheses. Describe your findings using simple, non-technical language. What does your analysis tell you about murder rates after the death penalty was suspended? What can you claim about causation from your statistical analysis?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**You're done! Congratulations.** 
Be sure to...

- run the tests and verify that they all pass, 
- choose **Download as PDF** from the **File** menu
- submit the .pdf file on **canvas**.

Tests for q1_1 (points: None)
```
>>> # TEST
>>> len(ab_test_order) == 6
True
```
```
>>> # TEST
>>> correct_order = make_array(5, 1, 3, 2, 4, 6)
>>> all(correct_order == ab_test_order)
True
```


Tests for q2_1 (points: None)
```
>>> # TEST
>>> ak_mn.num_rows == 44
True
```
```
>>> # TEST
>>> ak_mn.column("Murder rate in Alaska").item(0) == 10.19999981
True
```
```
>>> # TEST
>>> ak_mn.column("Murder rate in Minnesota").item(0) == 1.200000048
True
```


Tests for q3_10 (points: None)
```
>>> # TEST
>>> test_stat = round(simulate_and_test_statistic(change_in_death_rates, "Death Penalty", "Murder Rate"), 3)
>>> -5 < test_stat < 5
True
```


Tests for q3_11 (points: None)
```
>>> # TEST
>>> len(differences) == 5000
True
```
```
>>> # TEST
>>> abs(np.average(differences)) < 1
True
```
```
>>> # TEST
>>> all(differences == differences.item(0)) == False
True
```


Tests for q3_12 (points: None)
```
>>> # TEST
>>> empirical_P > 0.05
True
```


Tests for q3_3 (points: None)
```
>>> # TEST
>>> isinstance(preban_rates, Table)
True
```
```
>>> # TEST
>>> preban_rates.num_rows == 44
True
```
```
>>> # TEST
>>> np.all(preban_rates.column("Death Penalty") == True)
True
```
```
>>> # TEST
>>> np.all(preban_rates.column("Year") == 1971)
True
```
```
>>> # TEST
>>> all(elem in death_penalty.column("State") for elem in preban_rates.column("State"))
True
```


Tests for q3_4 (points: None)
```
>>> # TEST
>>> isinstance(postban_rates, Table)
True
```
```
>>> # TEST
>>> postban_rates.num_rows == 44
True
```
```
>>> # TEST
>>> np.all(postban_rates.column("Death Penalty") == False)
True
```
```
>>> # TEST
>>> np.all(postban_rates.column("Year") == 1973)
True
```
```
>>> # TEST
>>> all(elem in postban_rates.column("State") for elem in preban_rates.column("State"))
True
```


Tests for q3_6 (points: None)
```
>>> # TEST
>>> rate_means.num_rows == 2
True
```
```
>>> # TEST
>>> round(rate_means.where("Death Penalty", False).column(1).item(0), 15) == 8.120454540452272
True
```
```
>>> # TEST
>>> round(rate_means.where("Death Penalty", True).column(1).item(0), 15) == 7.513636380386362
True
```


Tests for q3_8 (points: None)
```
>>> # TEST
>>> isinstance(observed_difference, float)
True
```
```
>>> # TEST
>>> round(observed_difference, 3) == 0.607
True
```


Tests for q3_9 (points: None)
```
>>> # TEST
>>> np.isclose(round(find_test_stat(change_in_death_rates, "Death Penalty", "Murder Rate"),3) - 0.607, 0)
True
```
